```python
# Initialize packages
import pandas as pd
import altair as alt
import numpy as np
```

# Project Background

This project aims to build a model to use energy consumptions per capita to predict CO2 Emission per capita of a country.

<!-- BEGIN QUESTION -->

<div class="alert alert-info" style="color:black">

## Data Cleaning


<p>The data cleaning part includes:</p>

<ul>
  <li>Importing Data.</li>
    <ul>
      <li><a href="https://www.gapminder.org/free-material/">
          FREE DATA FROM WORLD BANK VIA GAPMINDER.ORG.</a></li>
    </ul>
  <li>Data cleaning and data wrangling using <code>pandas</code> and <code>altair</code>.</li>
</ul>
</div>

```python
# read datasets and melt dataframe from wide table to long table
# co2 emission per capita
co2_e = pd.read_csv("Data/Raw/co2_emissions_tonnes_per_person.csv").melt(id_vars='country', var_name='year', value_name='co2_e')
# coal consumption per capita
coal_c = pd.read_csv("Data/Raw/coal_consumption_per_cap.csv").melt(id_vars='country', var_name='year', value_name='coal_c')
# electricity generation per capita
elec_g = pd.read_csv("Data/Raw/electricity_generation_per_person.csv").melt(id_vars='country', var_name='year', value_name='elec_g')
# electricity consumption per capita
elec_c = pd.read_csv("Data/Raw/electricity_use_per_person.csv").melt(id_vars='country', var_name='year', value_name='elec_c')
# hydro generation per capita
hydro_g = pd.read_csv("Data/Raw/hydro_power_generation_per_person.csv").melt(id_vars='country', var_name='year', value_name='hydro_g')
# nuclear generation per capita
nuclear_g = pd.read_csv("Data/Raw/nuclear_power_generation_per_person.csv").melt(id_vars='country', var_name='year', value_name='nuclear_g')
# gas generation per capita
gas_g = pd.read_csv("Data/Raw/natural_gas_production_per_person.csv").melt(id_vars='country', var_name='year', value_name='gas_g')
# oil consumption per capita
oil_c = pd.read_csv("Data/Raw/oil_consumption_per_cap.csv").melt(id_vars='country', var_name='year', value_name='oil_c')
# oil generation per capita
oil_g = pd.read_csv("Data/Raw/oil_production_per_person.csv").melt(id_vars='country', var_name='year', value_name='oil_g')
coal_c
```

Output:
```
country  year   coal_c
0              UAE  1965      0.0
1        Argentina  1965   0.0328
2        Australia  1965     1.54
3          Austria  1965    0.696
4       Azerbaijan  1965      NaN
...            ...   ...      ...
4340           USA  2019    0.823
4341    Uzbekistan  2019   0.0526
4342     Venezuela  2019  0.00265
4343       Vietnam  2019    0.513
4344  South Africa  2019     1.56

[4345 rows x 3 columns]
```

```python
#merging data
dataframes = [co2_e, coal_c, elec_g, elec_c, hydro_g, nuclear_g, gas_g, oil_c, oil_g]
merged_df = co2_e
for df in dataframes[1:]:
    # Merge each DataFrame with the merged DataFrame
    merged_df = pd.merge(merged_df, df, on=['country','year'], how ='outer')
merged_df
```

Output:
```
country  year co2_e coal_c elec_g elec_c hydro_g nuclear_g  \
0            Afghanistan  1800   NaN    NaN    NaN    NaN     NaN       NaN   
1                 Angola  1800   NaN    NaN    NaN    NaN     NaN       NaN   
2                Albania  1800   NaN    NaN    NaN    NaN     NaN       NaN   
3                Andorra  1800   NaN    NaN    NaN    NaN     NaN       NaN   
4                    UAE  1800   NaN    NaN    NaN    NaN     NaN       NaN   
...                  ...   ...   ...    ...    ...    ...     ...       ...   
42803              Yemen  2019   NaN    NaN    NaN    NaN     NaN       NaN   
42804             Zambia  2019   NaN    NaN    NaN    NaN     NaN       NaN   
42805           Zimbabwe  2019   NaN    NaN    NaN    NaN     NaN       NaN   
42806  Equatorial Guinea  2019   NaN    NaN    NaN    NaN     NaN       NaN   
42807               Chad  2019   NaN    NaN    NaN    NaN     NaN       NaN   

       gas_g  oil_c  oil_g  
0        NaN    NaN    NaN  
1        N
```

```python
#EDA to remove years and countries with too many NaN
cleaned_df = merged_df.query('not co2_e.isna() and not coal_c.isna() and not elec_c.isna() and not oil_c.isna()')
# remove rows that co2_e, coal_c, elec_c and oil_c do not have na value. But we think it's acceptable that elec_g, hydro_g, nuclear_g and gas_g have na value. 
# Filling NaN with 0.
cleaned_df = cleaned_df.fillna(0)
cleaned_df = cleaned_df.reset_index().drop(columns='index')
cleaned_df
```

Output:
```
country  year  co2_e coal_c elec_g  elec_c  hydro_g nuclear_g  \
0        Australia  1965  10.70   1.54      0    2630   0.0628         0   
1          Austria  1965   5.23  0.696      0    2310    0.186         0   
2          Belgium  1965  11.20   2.03      0    2160  0.00248         0   
3           Canada  1965  12.80  0.788      0    6910    0.516   0.00056   
4      Switzerland  1965   5.22  0.219      0    3590    0.364         0   
...            ...   ...    ...    ...    ...     ...      ...       ...   
3340  South Africa  2014   8.86   1.64   4670    4180        0         0   
3341     Indonesia  2015   1.96  0.198    906   910.0        0         0   
3342     Indonesia  2016   2.15  0.204    948   956.0        0         0   
3343     Indonesia  2017   2.21  0.216    962  1020.0        0         0   
3344     Indonesia  2018   2.30  0.253    998  1060.0        0         0   

      gas_g  oil_c   oil_g  
0     0.000  1.340  0.0306  
1     0.000  0.761  0.0000  
2     0.000
```

```python
# save dataframe
cleaned_df.to_csv('Data/Processed/save_the_earth_processed_data.csv', index=True)
```

<div class="alert alert-danger" style="color:black">
    
**Credential**
    
FREE DATA FROM WORLD BANK VIA GAPMINDER.ORG, CC-BY LICENSE
</div>

Tests for q10 (points: None)
```
>>> assert area_co2 is not None, 'Your answer does not exist. Have you passed in the correct variable?'
>>> assert isinstance(area_co2, alt.Chart), 'Your answer is not an altair Chart object. Check to make sure that you have assigned an alt.Chart object.'
```


Tests for q11 (points: 5)
```
>>> assert income_hist is not None, 'Your answer does not exist. Have you passed in the correct variable?'
>>> assert isinstance(income_hist, alt.FacetChart), 'Your answer is not an altair Chart object. Check to make sure that you have assigned an alt.Chart object.'
```


Tests for q12c (points: None)
```
>>> assert scatter_lifeexp_familysize_styled is not None, 'Your answer does not exist. Have you passed in the correct variable?'
>>> assert isinstance(scatter_lifeexp_familysize_styled, alt.Chart), 'Your answer is not an altair Chart object. Check to make sure that you have assigned an alt.Chart object.'
```


Tests for q1w (points: 3)
```
>>> assert scatter is not None
```
failure_message: Did you assign the plot to the correct variable?
```
>>> assert isinstance(scatter, alt.Chart)
```
failure_message: Your answer is not an altair Chart object.
```
>>> spec = scatter.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['mark']['type'] == 'point'
```
failure_message: Make sure you are using the correct mark type.
```
>>> spec = scatter.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['x']['field'] == 'Weight_in_lbs'
```
failure_message: The wrong dataframe column is used for your chart's x channel.
```
>>> spec = scatter.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['x']['type'] == 'quantitative'
```
failure_message: The wrong data type is used for your chart's x channel.
```
>>> spec = scatter.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['y']['field'] == 'Horsepower'
```
failure_message: The wrong dataframe column is used for your chart's y channel.
```
>>> spec = scatter.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['y']['type'] == 'quantitative'
```
failure_message: The wrong data type is used for your chart's y channel.
```
>>> spec = scatter.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['color']['field'] == 'Origin'
```
failure_message: The wrong dataframe column is used for your chart's color channel.
```
>>> spec = scatter.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['color']['type'] == 'nominal'
```
failure_message: The wrong data type is used for your chart's color channel.


Tests for q2w (points: 3)
```
>>> assert price_lineplot is not None
```
failure_message: Did you assign the plot to the correct variable?
```
>>> assert isinstance(price_lineplot, alt.Chart)
```
failure_message: Your answer is not an altair Chart object.
```
>>> spec = price_lineplot.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['mark']['type'] == 'line'
```
failure_message: Make sure you are using the correct mark type.
```
>>> spec = price_lineplot.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['x']['field'] == 'date'
```
failure_message: The wrong dataframe column is used for your chart's x channel.
```
>>> spec = price_lineplot.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['x']['type'] == 'temporal'
```
failure_message: The wrong data type is used for your chart's x channel.
```
>>> spec = price_lineplot.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['y']['field'] == 'price'
```
failure_message: The wrong dataframe column is used for your chart's y channel.
```
>>> spec = price_lineplot.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['y']['type'] == 'quantitative'
```
failure_message: The wrong data type is used for your chart's y channel.
```
>>> spec = price_lineplot.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['color']['field'] == 'symbol'
```
failure_message: The wrong dataframe column is used for your chart's color channel.
```
>>> spec = price_lineplot.copy()
>>> with alt.data_transformers.enable('default'):
...     spec['data'] = spec['data'][:1]
...     spec = spec.to_dict()
>>> assert spec['encoding']['color']['type'] == 'nominal'
```
failure_message: The wrong data type is used for your chart's color channel.


Tests for q4 (points: None)
```
>>> assert scatter_familysize_lifeexp is not None, 'Your answer does not exist. Have you passed in the correct variable?'
>>> assert isinstance(scatter_familysize_lifeexp, alt.Chart), 'Your answer is not an altair Chart object. Check to make sure that you have assigned an alt.Chart object.'
```


Tests for q5 (points: None)
```
>>> assert line_points_education_ratio is not None, 'Your answer does not exist. Have you passed in the correct variable?'
>>> assert isinstance(line_points_education_ratio, alt.LayerChart), 'Your answer is not a layered altair Chart object. Check to make sure that you have assigned an alt.Chart object.'
```


Tests for q6c (points: None)
```
>>> assert ci_bands_education_ratio is not None, 'Your answer does not exist. Have you passed in the correct variable?'
>>> assert isinstance(ci_bands_education_ratio, alt.LayerChart), 'Your answer is not a layered altair Chart object. Check to make sure that you have assigned an alt.Chart object.'
```


Tests for q7 (points: None)
```
>>> assert scatter_familysize_mortality is not None, 'Your answer does not exist. Have you passed in the correct variable?'
>>> assert type(scatter_familysize_mortality) == alt.FacetChart, 'Your answer is not an altair Facet Chart object. Check to make sure that you have assigned an alt.Chart object.'
```


Tests for q9 (points: None)
```
>>> assert bars_co2 is not None, 'Your answer does not exist. Have you passed in the correct variable?'
>>> assert isinstance(bars_co2, alt.Chart), 'Your answer is not an altair Chart object. Check to make sure that you have assigned an alt.Chart object.'
```
